<div class = "alert alert-warning" align="center"> 
    <img src="https://www.egd.mg/wp-content/uploads/2020/01/cropped-logo-home-110px.png" align="left" width="290"  margin="50px"> 
    <center><h1>Lycée français de Tananarive</h1> 
        <h2> Jupyter Hub,  Spécialité NSI </h2></center>
</div>

### 3. Ecrire une fonction logique à partir de son expression.
Considérons la fonction logique décrite par l'expression ci dessous :
$$S = (\bar{A} + B) \times C $$

Ecrire une fonction `logique1(A,B,C)` qui renvoie le résultat booléen de cette expression.

<!-- grade_id: cell-137d01c21c7c4f91, points: 0, grade: false, solution: true -->
```python

```

Output:
```
False
```

Vérifions maintenant que votre fonction passe bien le jeu de tests :

<!-- grade_id: cell-cc6fd57a98b27ab4, points: 2, grade: true, solution: false -->
```python
assert logique1(False,False,False) == False
assert logique1(False,True,True) == True
assert logique1(False,False,True) == True
assert logique1(True,True,False) == False
```

### 4. Ecrire une fonction logique à partir de sa table de vérité.
Considérons la table de vérité ci dessous :

|  A  |  B  |  C  |  S  |
|:---:|:---:|:---:|:---:|
|0|0|0|1|
|0|0|1|0|
|0|1|0|1|
|0|1|1|0|
|1|0|0|1|
|1|0|1|0|
|1|1|0|1|
|1|1|1|1|

Ecrire une fonction `logique2(A,B,C)` qui renvoie une valeur qui corresponde à cette table de vérité.

<!-- grade_id: cell-1cd5540bfe60d066, points: 0, grade: false, solution: true -->
```python

```

Output:
```
True
```

Quelques tests qui vont vérifier la conformité de votre fonction.

<!-- grade_id: cell-2782b2aa2a543221, points: 3, grade: true, solution: false -->
```python
assert logique2(False,False,False) == True
assert logique2(False,True,True) == False
assert logique2(True,True,False) == True
assert logique2(True,True,True) == True
```

### 5. Ecrire une fonction logique à partir d'un circuit logique
Considérons le circuit ci-dessous
![20% center](circuit.png)
Ecrire une fonction `logique3(A,B,C)` qui renvoie une valeur qui corresponde à ce circuit.

<!-- grade_id: cell-d726d50776cc2830, points: 0, grade: false, solution: true -->
```python

```

Output:
```
True
```

Quelques tests qui vont vérifier la conformité de votre fonction.

<!-- grade_id: cell-c5a80c7c71ebac29, points: 2, grade: true, solution: false -->
```python
assert logique3(False,False,False) == True
assert logique3(False,True,True) == True
assert logique3(True,False,True) == True
assert logique3(True,True,True) == False
```

<div class="alert alert-warning" align="center">
    <b>Equipe pédagoqique LFT-NSI</b>
<br/>ressource éducative libre distribuée sous <a href="http://creativecommons.org/licenses/by-nc-sa/4.0/"!>Licence Creative Commons Attribution</a>
    <img src="https://i.creativecommons.org/l/by-nc-sa/4.0/88x31.png" align="center" width="70"  margin="50px"> 
</div>